## Lab02: Logic

- Full Name:

- Student ID:

<font color='red'>**Attention:**</font>

You can discuss ideas with classmates as well as finding information from the internet, book, etc...; but *this homework must be your*.

**How to submit your homework**

Before submitting, rerun the notebook (`Kernel` ->` Restart & Run All`).

Create folder named `Student_ID` (for example, if your student id is 1234567, you will name the folder as `1234567`) with the following structure:

   - Folder `Project02_logic`
        - File `Project02_logic.ipynb` 
        - Folder `models` and file `logic.py`
        - Folder `PS4` that include your works on exercise 4:
            - `SRC`: Source code
            - `REPORT`: Report (file .pdf)

Compress `Student_ID` folder with <font color='red'>**zip**</font> format.

<font color=red>Please, follow exactly all the requirements and instructions above.</font>



## How to do your homework: 

Contents:
- Propositional logic

For questions 1-3, you will work directly on this notebook. The sentence ```# YOUR CODE HERE``` will indicate the parts you need to do. Mathematical symbols must be presented with code as in the following symbol list: 

<center>
    
|  | Math symbols | Code | 
|:---|:----:|---|
| Ký hiệu hằng số (constant symbol) | hcmus | ```Constant('hcmus')``` (ghi thường) | 
| Ký hiệu biến (variable symbol) | $x$ | ```Variable('$x')``` (ghi thường) | 
| Câu nguyên tố (atomic formula) | Rain, LocatedIn(hcmus, $x$) | ```Atom('Rain')```, ```Atom('LocatedIn', 'hcmus', '$x')``` <br> (Kí hiệu vị từ ghi hoa, đối số ghi thường) |
| Phép phủ định (negation) | $\neg$ Rain | ```Not(Atom('Rain'))```|
| Phép và (conjunction)| Rain $\land$ Snow | ```And(Atom('Rain'), Atom('Snow'))```|
| Phép hoặc (disjunction) | Rain $\lor$ Snow | ```Or(Atom('Rain'), Atom('Snow'))```|
| Phép kéo theo (implication) | Rain $\Rightarrow$ Wet | ```Implies(Atom('Rain'), Atom('Wet'))```|
| Phép tương đương (equivalence) | Rain $\Leftrightarrow$ Wet | ```Equiv(Atom('Rain'), Atom('Wet'))```|
| Lượng từ tồn tại (existential quantification) | $\exists x$ LocatedIn(hcmus,$x$) | ```Exists('$x', Atom('LocatedIn', 'hcmus', '$x'))```|
| Lượng từ với mọi (universal quantification) | $\forall x$ MadeOfAtoms($x$) | ```Forall('$x', Atom('MadeOfAtoms', '$x'))```|    
| Phép và với nhiều đối số <br>(tương tự với phép hoặc)| A $\land$ B $\land$ C | ```AndList([Atom('A'), Atom('B'), Atom('C')]) ```|
  
    
</center>

```python
import pickle, gzip, os, random
import sys
import os
import collections

from logic import *
from typing import Tuple, List

def checkFormula(name, predForm, preconditionForm=None):
    filename = os.path.join('models', name + '.pklz')
    objects, targetModels = pickle.load(gzip.open(filename)) 
    
    # If preconditionion exists, change the formula to
    preconditionPredForm = And(preconditionForm, predForm) if preconditionForm else predForm
    predModels = performModelChecking([preconditionPredForm], findAll=True, objects=objects)
    ok = True
    
    def hashkey(model): return tuple(sorted(str(atom) for atom in model))
    targetModelSet = set(hashkey(model) for model in targetModels)
    predModelSet = set(hashkey(model) for model in predModels)
    for model in targetModels:
        if hashkey(model) not in predModelSet:
            print("Your formula (%s) says the following model is FALSE, but it should be TRUE:" % predForm)
            ok = False
            printModel(model)
            return
    for model in predModels:
        if hashkey(model) not in targetModelSet:
            print("Your formula (%s) says gaugau" % predForm)
            printModel(model)
            return
        
    print(f"You matched the {len(targetModels)} models")
    print(f"Example model: {rstr(random.choice(targetModels))}")
    return True

```

```python
# import other modules if you need them 

```

## 1. Propositional Logic (0.75đ)

For example: "If it 's rain, the road will be wet":

```
def rainWet():
    Rain = Atom('Rain') 
    Wet = Atom('Wet')
    return Implies(Rain, Wet)
```


**1a) If it's summer in California, it will not rain**

<!-- grade_id: TODO_1a, points: 0, grade: false, solution: true -->
```python
def formula1a() -> Formula:
    # Predicates to use:
    Summer = Atom('Summer')               # whether it's summer
    California = Atom('California')       # whether we're in California
    Rain = Atom('Rain')                   # whether it's raining
    
    # YOUR CODE HERE
    return Implies(And(Summer, California), Not(Rain))
```

<!-- grade_id: GRADE_1a, points: 0.25, grade: true, solution: false -->
```python
assert checkFormula('1a', formula1a()) == True
```

Output:
```
You matched the 7 models
Example model: {'Summer', 'California'}
```

**1b) The road is wet if and only if the weather is rainy or the sprinkler is opened**

<!-- grade_id: TODO_1b, points: 0, grade: false, solution: true -->
```python
def formula1b() -> Formula:
    # Predicates to use:
    Rain = Atom('Rain')              # whether it is raining
    Wet = Atom('Wet')                # whether it it wet
    Sprinklers = Atom('Sprinklers')  # whether the sprinklers are on
    
    # YOUR CODE HERE
    return Equiv(Wet,Or(Rain, Sprinklers))
```

<!-- grade_id: GRADE_1b, points: 0.25, grade: true, solution: false -->
```python
assert checkFormula('1b', formula1b()) == True
```

Output:
```
You matched the 4 models
Example model: {'Wet', 'Rain'}
```

**1c) Day or night**

<!-- grade_id: TODO_1c, points: 0, grade: false, solution: true -->
```python
def formula1c() -> Formula:
    # Predicates to use:
    Day = Atom('Day')     # whether it's day
    Night = Atom('Night') # whether it's night
    
    # YOUR CODE HERE
    return Or(And(Day,Not(Night)), And(Night,Not(Day)))

```

<!-- grade_id: GRADE_1c, points: 0.25, grade: true, solution: false -->
```python
assert checkFormula('1c', formula1c()) == True
```

Output:
```
You matched the 2 models
Example model: {'Day'}
```

## 2 First-order logic (2đ)

Formulate statement from natural languages to First-order logic.

For example: "There is a lighting source that shines"

```
def lightShines():
    def Light(x): return Atom('Light', x)    # 
    def Shines(x): return Atom('Shines', x)  # 
    return Exists('$x', And(Light('$x'), Shines('$x')))
```

**2a) Everybody has mother**

<!-- grade_id: TODO_2a, points: 0, grade: false, solution: true -->
```python
def formula2a() -> Formula:
    # Predicates to use:
    def Person(x): return Atom('Person', x)        # whether x is a person
    def Mother(x, y): return Atom('Mother', x, y)  # whether x's mother is y

    # Note: You do NOT have to enforce that the mother is a "person"
    # YOUR CODE HERE
    return Forall('$x', Implies(Person('$x'), Exists('$y',Mother('$x','$y'))))
    
```

<!-- grade_id: GRADE_2a, points: 0.5, grade: true, solution: false -->
```python
formula2a_precondition = AntiReflexive('Mother')
assert checkFormula('2a', formula2a(), formula2a_precondition) == True
```

Output:
```
You matched the 343 models
Example model: {'Mother(o1,o3)', 'Person(o3)', 'Mother(o3,o1)', 'Mother(o3,o2)'}
```

**2b) At least one person does not have any child**

<!-- grade_id: TODO_2b, points: 0, grade: false, solution: true -->
```python
def formula2b() -> Formula:
    # Predicates to use:
    def Person(x): return Atom('Person', x)        # whether x is a person
    def Child(x, y): return Atom('Child', x, y)    # whether x has a child y

    # Note: You do NOT have to enforce that the child is a "person"
    # YOUR CODE HERE
    return Exists('$x',And(Person('$x'), Forall('$y', Not(Child('$x', '$y')))))
```

<!-- grade_id: GRADE_2b, points: 0.5, grade: true, solution: false -->
```python
formula2b_precondition = AntiReflexive('Child')
assert checkFormula('2b', formula2b(), formula2b_precondition) == True
```

Output:
```
You matched the 169 models
Example model: {'Person(o1)', 'Person(o2)', 'Child(o3,o1)', 'Child(o3,o2)'}
```

**2c) Create definition for Daughter by Female and Child**

Examples

```
# Defnition of Parent by Child
def parentChild():
    def Parent(x, y): return Atom('Parent', x, y)  # whether x has a parent y
    def Child(x, y): return Atom('Child', x, y)    # whether x has a child y
    return Forall('$x', Forall('$y', Equiv(Parent('$x', '$y'), Child('$y', '$x'))))
    
```

<!-- grade_id: TODO_2c, points: 0, grade: false, solution: true -->
```python
def formula2c() -> Formula:
    # Predicates to use:
    def Female(x): return Atom('Female', x)            # whether x is female
    def Child(x, y): return Atom('Child', x, y)        # whether x has a child y
    def Daughter(x, y): return Atom('Daughter', x, y) # whether x has a daughter y
    
    # YOUR CODE HERE
    return Forall('$x', Forall('$y', Equiv(Daughter('$x', '$y'), And(Female('$y'), Child('$x', '$y')))))
```

<!-- grade_id: GRADE_2c, points: 0.5, grade: true, solution: false -->
```python
formula2c_precondition = AntiReflexive('Child')
assert checkFormula('2c', formula2c(), formula2c_precondition) == True
```

Output:
```
You matched the 512 models
Example model: {'Child(o2,o1)', 'Child(o2,o3)', 'Child(o1,o2)', 'Daughter(o1,o2)', 'Child(o3,o1)', 'Female(o2)'}
```

**2d) Create definition for Grandmother by Female and Parent**

<!-- grade_id: TODO_2d, points: 0, grade: false, solution: true -->
```python
# Note: It is ok for a person to be her own parent
def formula2d() -> Formula:
    # Predicates to use:
    def Female(x): return Atom('Female', x)                  # whether x is female
    def Parent(x, y): return Atom('Parent', x, y)            # whether x has a parent y
    def Grandmother(x, y): return Atom('Grandmother', x, y)  # whether x has a grandmother y
    
    # YOUR CODE HERE
    return Forall('$x', 
                  Forall('$y', Equiv(Grandmother('$x', '$y'), 
                                     And(Female('$y'), 
                                         Exists('$z', 
                                                And(Parent('$x', '$z'),Parent('$z', '$y'))
                                                )
                                         )
                                    )
                        )
                )
```

<!-- grade_id: GRADE_2d, points: 0.5, grade: true, solution: false -->
```python
formula2d_precondition = AntiReflexive('Parent')
assert checkFormula('2d', formula2d(), formula2d_precondition) == True
```

Output:
```
You matched the 512 models
Example model: {'Parent(o2,o3)', 'Female(o1)', 'Parent(o3,o1)', 'Female(o3)', 'Grandmother(o2,o1)'}
```

## 3. Who is the liar? (2.25đ)

Using prepositional method for the following problem:

A Facebook employee make a serious mistake that crashed the server, people are discussing to find out the one that caused this problem. After a while, four suspects are identified.

We have the following information:
- John: "I am not the guy that make the server down!"
- Susan: "It should be Nicole."
- Mark: "No, it was Susan!"
- Nicole: "Susan is lying, you should not trust her."
- There is **only 1 person** told the truth.
- There is **only 1 person** crashed the server.

Finish the ```liar()``` function with 6 formulas for finding the liar!

<!-- grade_id: TODO_3, points: 0, grade: false, solution: true -->
```python
# Problem 3: Liar puzzle

# Facts:
# 0. John: "It wasn't me!"
# 1. Susan: "It was Nicole!"
# 2. Mark: "No, it was Susan!"
# 3. Nicole: "Susan's a liar."
# 4. Exactly one person is telling the truth.
# 5. Exactly one person crashed the server.
# Query: Who did it?

# This function returns a list of 6 formulas corresponding to each of the
# above facts. Be sure your formulas are exactly in the order specified. 


# Hint: For fact 4 & 5, you might want to use the Equals predicate, defined as:
# def Equals(x, y): return Atom('Equals', x, y)
# This predicate is used to assert that two objects are the same.
# In particular, Equals(x,x) = True and Equals(x,y) = False iff x is not equal to y. 
# It can also be solved in other ways, without the Equals predicate!
    
def liar() -> Tuple[List[Formula], Formula]:
    def TellTruth(x): return Atom('TellTruth', x)
    def CrashedServer(x): return Atom('CrashedServer', x)
    def Equals(x, y): return Atom('Equals', x, y)
    john = Constant('john')
    susan = Constant('susan')
    nicole = Constant('nicole')
    mark = Constant('mark')
    formulas = []
    # We provide the formula for fact 0 here.
    formulas.append(Equiv(TellTruth(john), Not(CrashedServer(john))))
    
    # You should add 5 formulas, one for each of facts 1-5.
    # YOUR CODE HERE
    formulas.append(Equiv(TellTruth(susan), CrashedServer(nicole)))
    formulas.append(Equiv(TellTruth(mark), CrashedServer(susan)))
    formulas.append(Equiv(TellTruth(nicole), And(Not(TellTruth(susan)),Not(CrashedServer(nicole)))))
    formulas.append(
        Implies(
            And(
                Forall('$x', TellTruth('$x')), 
                Forall('$y', Not(Equals('$x', '$y')))
            ),
            Not(TellTruth('$y'))
        )
    )
    
    formulas.append(
        Implies(
            And(
                Forall('$x', CrashedServer('$x')), 
                Forall('$y', Not(Equals('$x', '$y')))
            ), 
            Not(CrashedServer('$y'))
        )
    )
    
    query = CrashedServer('$x')
    return (formulas, query)
```

<!-- grade_id: GRADE_3_len, points: 0.25, grade: true, solution: false -->
```python
predForms, predQuery = liar()
assert len(predForms) == 6
```

```python
predForms, predQuery = liar()
formula_id = 0
assert checkFormula('3a-' + str(formula_id), predForms[formula_id]) == True
```

Output:
```
You matched the 2 models
Example model: {'TellTruth(john)'}
```

<!-- grade_id: GRADE_3_1, points: 0.25, grade: true, solution: false -->
```python
predForms, predQuery = liar()
formula_id = 1
assert checkFormula('3a-' + str(formula_id), predForms[formula_id]) == True
```

Output:
```
You matched the 2 models
Example model: set()
```

<!-- grade_id: GRADE_3_2, points: 0.25, grade: true, solution: false -->
```python
predForms, predQuery = liar()
formula_id = 2
assert checkFormula('3a-' + str(formula_id), predForms[formula_id]) == True
```

Output:
```
You matched the 2 models
Example model: {'CrashedServer(susan)', 'TellTruth(mark)'}
```

<!-- grade_id: GRADE_3_3, points: 0.25, grade: true, solution: false -->
```python
predForms, predQuery = liar()
formula_id = 3
assert checkFormula('3a-' + str(formula_id), predForms[formula_id]) == True
```

Output:
```
Your formula (And(Implies(TellTruth(nicole),And(Not(TellTruth(susan)),Not(CrashedServer(nicole)))),Implies(And(Not(TellTruth(susan)),Not(CrashedServer(nicole))),TellTruth(nicole)))) says gaugau
('*', 'CrashedServer(nicole)', '=', 'True')
('*', '(other atoms if any)', '=', 'False')
AssertionError                            Traceback (most recent call last)
Cell In[69], line 3
      1 predForms, predQuery = liar()
      2 formula_id = 3
----> 3 assert checkFormula('3a-' + str(formula_id), predForms[formula_id]) == True

AssertionError:
```

<!-- grade_id: GRADE_3_4, points: 0.5, grade: true, solution: false -->
```python
predForms, predQuery = liar()
formula_id = 4
assert checkFormula('3a-' + str(formula_id), predForms[formula_id]) == True
```

Output:
```
File ~/Downloads/MSc_S4_AppliedAIAndAlgorithm/Labs/Lab02/Student_ID/Project02_logic/logic.py:622, in <listcomp>(.0)
    620 if form.isa(Constant): return form
    621 if form.isa(Atom):
--> 622     return Atom(*[form.name] + [convert(arg, subst) for arg in form.args])
    623 if form.isa(Not): return Not(convert(form.arg, subst))
    624 if form.isa(And): return And(convert(form.arg1, subst), convert(form.arg2, subst))

File ~/Downloads/MSc_S4_AppliedAIAndAlgorithm/Labs/Lab02/Student_ID/Project02_logic/logic.py:618, in propositionalize.<locals>.convert(form, subst)
    616 def convert(form, subst):
    617     if form.isa(Variable):
--> 618         if form not in subst: raise Exception("Free variable found: %s" % form)
    619         return subst[form]
    620     if form.isa(Constant): return form

Exception: Free variable found: $x
```

<!-- grade_id: GRADE_3_5, points: 0.5, grade: true, solution: false -->
```python
predForms, predQuery = liar()
formula_id = 5
assert checkFormula('3a-' + str(formula_id), predForms[formula_id]) == True
```

Output:
```
File ~/Downloads/MSc_S4_AppliedAIAndAlgorithm/Labs/Lab02/Student_ID/Project02_logic/logic.py:622, in <listcomp>(.0)
    620 if form.isa(Constant): return form
    621 if form.isa(Atom):
--> 622     return Atom(*[form.name] + [convert(arg, subst) for arg in form.args])
    623 if form.isa(Not): return Not(convert(form.arg, subst))
    624 if form.isa(And): return And(convert(form.arg1, subst), convert(form.arg2, subst))

File ~/Downloads/MSc_S4_AppliedAIAndAlgorithm/Labs/Lab02/Student_ID/Project02_logic/logic.py:618, in propositionalize.<locals>.convert(form, subst)
    616 def convert(form, subst):
    617     if form.isa(Variable):
--> 618         if form not in subst: raise Exception("Free variable found: %s" % form)
    619         return subst[form]
    620     if form.isa(Constant): return form

Exception: Free variable found: $x
```

<!-- grade_id: GRADE_3_ALL, points: 0.25, grade: true, solution: false -->
```python
predForms, predQuery = liar()
assert checkFormula('3a-all', AndList(predForms)) == True
```

Output:
```
File ~/Downloads/MSc_S4_AppliedAIAndAlgorithm/Labs/Lab02/Student_ID/Project02_logic/logic.py:622, in <listcomp>(.0)
    620 if form.isa(Constant): return form
    621 if form.isa(Atom):
--> 622     return Atom(*[form.name] + [convert(arg, subst) for arg in form.args])
    623 if form.isa(Not): return Not(convert(form.arg, subst))
    624 if form.isa(And): return And(convert(form.arg1, subst), convert(form.arg2, subst))

File ~/Downloads/MSc_S4_AppliedAIAndAlgorithm/Labs/Lab02/Student_ID/Project02_logic/logic.py:618, in propositionalize.<locals>.convert(form, subst)
    616 def convert(form, subst):
    617     if form.isa(Variable):
--> 618         if form not in subst: raise Exception("Free variable found: %s" % form)
    619         return subst[form]
    620     if form.isa(Constant): return form

Exception: Free variable found: $x
```

Hãy chạy cell dưới đây để xem bạn đã giải được bài toán đặt ra chưa nhé :) 

```python
# Run this cell to solve the puzzle and find out who was the liar!
predForms, predQuery = liar()
kb = createModelCheckingKB()

filename = os.path.join('models', '3a-all.pklz')
objects, targetModels = pickle.load(gzip.open(filename))
for obj in objects:
    kb.tell(Atom('Object', obj))

for predForm in predForms:
    response = kb.tell(predForm)

response = kb.ask(predQuery)
showKBResponse(response)
```

Output:
```
File ~/Downloads/MSc_S4_AppliedAIAndAlgorithm/Labs/Lab02/Student_ID/Project02_logic/logic.py:795, in KnowledgeBase.tell(self, form)
    794 def tell(self, form):
--> 795     return self.query(form, modify=True)

File ~/Downloads/MSc_S4_AppliedAIAndAlgorithm/Labs/Lab02/Student_ID/Project02_logic/logic.py:816, in KnowledgeBase.query(self, form, modify)
    814 if len(freeVars) > 0:
    815     if modify:
--> 816         raise Exception("Can't modify database with a query with free variables: %s" % form)
    817     var = freeVars[0]
    818     allForms = AndList([deriv.form for deriv in list(self.derivations.values())])

Exception: Can't modify database with a query with free variables: Implies(And(Exists($x,TellTruth($x)),Forall($y,Not(Equals($x,$y)))),Not(TellTruth($y)))
```

## 4. Inference (5đ)

**You will designed this part by your own python scripts. Please do not write to this notebook**

Given a Knowledge Base **(KB)** and a statement $\alpha$, both represented by prepositional logic and change to CNF. Determine if KB entails $\alpha$ (KB ⊨ $\alpha$) by resolution.

- Your report have at least 5 test cases for demo (test cases should not be too simple). 
- Write short evaluations about the advantages and disadvantages of resolution method for prepositional logic, propose your own solution for specific problem.
**a. Desribe the experiment data.**
- **Input data**: **KB** and $\alpha$ with CNF stored in **input.txt**. This file have the following structure:
    - The first line contains $\alpha$ statement
    - The second line has N – the number of clauses in the KB.
    - N following lines are clauses in KB, one line for each clause.
    
Positive literal is represented by an uppercase character (A-Z). Negative literal is positive literal with minus (‘-‘) before the character.

Literals are linked by OR. Bettween literals and keywords there can be more than one space.

- **Output data**: The set of statements that are generated during the whole resolution process and the conclusion clause is stored in **output.txt**. This file must have the following structure:
    - The first line contains M1 – the number of generated clauses in first loop. M1 following lines describe those clauses (including empty clause), each line represents a clause. The empty clause is represented by “{}”
    - Next loops are also represented as ($M_2$, $M_3$,…, $M_n$ clauses) the first loop. 
    - The last line is the conclusion clause, which mean if “KB entails $\alpha$?”. Print YES if KB entails $\alpha$. NO for the opposite.
    - Remove redundant clause (same clauses in same loop, or the loops before).
    
    
- main function must::
    - Read the input data and store it in a suitable data structure
    - Call PL-Resolution function for execute resolution algorithm
    - Export the output data with the required format


- Store the semantic information of true and false value in PL-RESOLUTION. Do not forget that we need to negate $\alpha$.

- <font color='red'>Literals in a clause is sorted by alphabetical ordering </font>
- <font color='red'>Infered condition is checked in each loop as when all the new clauses are generated from the current KB.</font>
- <font color='red'>Clauses that have $A \vee B \vee -B$ format and $True$ value are similar to $A \vee True$ can be removed. </font>


Example: A KB and $\alpha$ clause in **input.txt**

| Input.txt | Output.txt | Ghi chú |
| --- | --- | --- |
| -A | 3 |  |
| 4 | -A | (-A OR B) hợp giải với (-B)|
| -A OR B | B | (-A OR B) hợp giải với (negative of -A) |
| B OR -C | -C | (-C OR B) hợp giải với (-B) |
| A OR -B OR C | 4 |  |
| -B | -B OR C | (A OR C OR -B) hợp giải với (-A) |
|  | A OR C | (A OR C OR -B) hợp giải với (B) |
|  | A OR -B| (A OR C OR -B) hợp giải với (-C) |
|  | {}| (-B) hợp giải với (B) |
|  | **YES** | KB entails $\alpha$ vì tồn tại mệnh đề rỗng trong KB |


Another example KB with a different $\alpha$

| Input.txt | Output.txt | Ghi chú |
| --- | --- | --- |
| A | 2 | KB KHÔNG entail $\alpha$ vì không phát sinh được mệnh đề mới và không tìm thấy mệnh đề rỗng|
| 4 | -C |  |
| -A OR B | -B OR C |  |
| -C OR B | 2 |  |
| A OR C OR -B | A OR -B |  |
| -B | -A OR C |  |
|  | 1 |  |
|  | A OR -C |  |
|  | **0** |  |
|  | NO |  |







**b. Evaluations**

| **STT** | **Description** | **Ratio** |
| --- | --- | --- |
| 1 | Read input data and store in suitable data structure | 0.5 |
| 2 | Implementation of resolution method  | 1.0 |
| 3 | Inference process and results | 2.5 |
| 5 | Testcases, report, evaluations,  | 1.0 |